# Численное Интегрирование (обязательно сделать хотя бы 1 из звёздочек)

# Квадратурные формулы Ньютона-Котеса.

Пусть некоторая функция $f(x)$ задана в узлах интерполяции $x_{i}=x_{0}+i h(i=1,2, \ldots, n)$ на отрезке $[a, b]$ таблицей своих значений:

\begin{array}{c|c|c|c|c}
x_{0}=a & x_{1} & x_{2} & \ldots & x_{n}=b \\
\hline y_{0}=f\left(x_{0}\right) & y_{1}=f\left(x_{1}\right) & y_{2}=f\left(x_{2}\right) & \ldots & y_{n}=f\left(x_{n}\right)
\end{array} 

По заданным значениям подынтегральной терполяционный многочлен Лагранжа. Для равноотстояших узлов интерноляционный многочлен Лагранжа. Для равноотстоящих узлов интерполяционный многочлен Лагранжа имеет вид  

$
L_{n}(x)=\sum_{i=0}^{n} \frac{(-1)^{n-1}}{i !(n-i) !} \frac{q(q-1) \ldots(q-n)}{q-1} y_{i}  
$

$
\text { где } q=\frac{x-x_{0}}{h}-\text { шаг интерполяции. } 
$

Заменим подынтегральную функцию $f(x)$ интерполяционным многочленом Лагранжа, тогда
$$
\int_{a}^{b} f(x) d x \approx \int_{x_{i}}^{x_{n}}\left[\sum_{i=0}^{n} \frac{(-1) n-1}{i !(n-i) !} \frac{q(q-1) \ldots(q-n)}{q-i} y_{i}\right] d x .
$$
Поменяем местами знак суммирования и интеграл и вынесем за знак интеграла постоянные коэффициенты:
$$
\int_{a}^{b} f(x) d x \approx \sum_{i=0}^{n} y_{i} \frac{(-1)^{n-1}}{i !(n-i) !} \int_{x_{0}}^{x_{n}} \frac{q(q-1) \ldots(q-n)}{q-i} d x .
$$
Поскольку $d q=\frac{d x}{h}$ и $h=\frac{b-a}{n}$, то, заменив пределы интегрирования, получим
$$
\begin{gathered}
\int_{a}^{b} f(x) d x \approx(b-a) \times \\
\times \sum_{i=0}^{n} y_{i} \frac{(-1)^{n-i}}{i !(n-i) !} \frac{1}{n} \int_{0}^{n} q(q-1) \ldots(q-i+1)(q-i-1) \ldots(q-n) d q .
\end{gathered}
$$

Положим

$$
H_{i}=\frac{(-1)^{n-i}}{i !(n-i) !} \frac{1}{n} \int_{0}^{n} q(q-1) \ldots(q-i+1)(q-i-1) \ldots(q-n) d q
$$

где $i=0,1,2, \ldots, n$.  

Числа $H_{i}$ называются коэффициентами Ньютона - Котеса. Они не зависят от вида функции $f(x)$ и являются функцией только $n$ (количества узлов интерполирования). Коэффициенты Ньютона - Котеса можно вычислить заранее для различного числа узлов интерполирования и свести в табл. 4.1.
Окончательный вид формулы Ньютон - Котеса:
$$
\int_{a}^{b} f(x) d x \approx(b-a) \sum_{i=0}^{n} y_{i} H_{i}
$$  

\begin{array}{|c|c|}
\hline n=1 & H_{0}=H_{1}=\frac{1}{2} \\
\hline n=2 & H_{0}=H_{2}=\frac{1}{6}, H_{1}=\frac{2}{3} \\
\hline n=3 & H_{0}=H_{3}=\frac{1}{8}, H_{1}=H_{2}=\frac{3}{8} \\
\hline n=4 & H_{0}=H_{4}=\frac{7}{90}, H_{1}=H_{3}=\frac{16}{45}, H_{2}=\frac{2}{15} \\
\hline n=5 & H_{0}=H_{5}=\frac{19}{288}, H_{1}=H_{4}=\frac{25}{96}, H_{2}=H_{3}=\frac{25}{144} \\
\hline n=6 & H_{0}=H_{6}=\frac{41}{840}, H_{1}=H_{5}=\frac{9}{35}, H_{2}=H_{4}=\frac{9}{280}, H_{3}=\frac{34}{105} \\
\hline n=7 & H_{0}=H_{7}=\frac{751}{17280}, H_{1}=H_{6}=\frac{3577}{17280}, H_{2}=H_{5}=\frac{1323}{17280}, H_{3}=H_{4}=\frac{2989}{17280} \\
\hline
\end{array}   

# Квадратурная формула Чебышева.

П. Л. Чебышев предложил для вычисления определенных интегралов воспользоваться формулой
$$
\int_{a}^{b} f(x) d x \approx(b-a) \sum_{i=0}^{n} c_{i} f\left(x_{i}\right),
$$
в которой квадратурные коэффициенты $c_{i}(i=1,2, \ldots, n)$ фиксированы, а абсциссы $x_{i}(i=1,2, \ldots, n)$ нужно определить.

Для упрощения вычислений выберем $c_{1}=c_{2}=\ldots=c_{n}$ и будем интегрировать на отрезке $[-1,1] .$ Тогда формула примет вид 
$$
\int_{-1}^{1} f(x) d x= {2*c}\left[f\left(x_{1}\right)+f\left(x_{2}\right)+\ldots+f\left(x_{n}\right)\right]  
$$

Коэффициент и узлы определим из условия, что это равенство является точным для случая, когда $f(x)-$ многочлен вида

$$
f(x)=a_{0}+a_{1} x+a_{2} x^{2}+a_{3} x^{3}+\ldots+a_{n} x^{n}
$$  

Подставим этот многочлен и проинтегрируем. В результате получим

$$
\int_{-1}^{1}\left(a_{0}+a_{1} x+a_{2} x^{2}+a_{3} x^{3}+\ldots+a_{n} x^{n}\right) d x=2\left(a_{0}+\frac{a_{2}}{3}+\frac{a_{4}}{5}+\ldots\right) .
$$

Далее


$f\left(x_{1}\right)=a_{0}+a_{1} x_{1}+a_{2} x_{1}^{2}+\ldots+a_{n} x_{1}^{n}$,
$f\left(x_{2}\right)=a_{0}+a_{1} x_{2}+a_{2} x_{2}^{2}+\ldots+a_{n} x_{2}^{n}$,
$f\left(x_{3}\right)=a_{0}+a_{1} x_{3}+a_{2} x_{3}^{2}+\ldots+a_{n} x_{n}^{n}$,
$f\left(x_{n}\right)=a_{0}+a_{1} x_{n}+a_{2} x_{n}^{2}+\ldots+a_{n} x_{n}^{n} .$


Тогда равенство  примет вид

$$
\begin{gathered}
2\left(a_{0}+\frac{a_{2}}{3}+\frac{a_{4}}{5}+\ldots\right)=2 c_{n}\left[n a_{0}+a_{1}\left(x_{1}+x_{2}+\ldots+x_{n}\right)+\right. \\
+a_{2}\left(x_{1}^{2}+x_{2}^{2}+\ldots+x_{n}^{2}\right)+a_{3}\left(x_{1}^{3}+x_{2}^{3}+\ldots+x_{n}^{3}\right)+\ldots \\
\left.\ldots+a_{n}\left(x_{1}^{n}+x_{2}^{n}+\ldots+x_{n}^{n}\right)\right] .
\end{gathered}
$$   





Полученное равенство должно выполняться при любых значениях $a_{0}, a_{1}, \ldots, a_{n} .$ Поэтому из сравнения коэффициентов $a_{i}$ в правой и левой частях соотношения (4) находим, что $n c_{n}=1$, откуда
$$
c_{n}=\frac{1}{n}
$$
и, кроме того,
$$
\left\{\begin{array}{c}
x_{1}+x_{2}+x_{3}+\ldots x_{n}=0 \\
x_{1}^{2}+x_{2}^{2}+x_{3}^{2}+\ldots+x_{n}^{2}=\frac{n}{3} \\
x_{1}^{3}+x_{2}^{3}+x_{3}^{3}+\ldots+x_{n}^{3}=0 \\
x_{1}^{4}+x_{2}^{4}+x_{3}^{4}+\ldots+x_{n}^{4}=\frac{n}{5} \\
\left.\cdot \cdot \cdot \cdot \cdot \cdot \cdot \cdot \cdot \cdot \cdot \cdot \cdot \cdot \cdot \cdot \cdot \cdot \cdot \cdot \cdot \cdot(-1)^{n+1}\right] \\
x_{1}^{n}+x_{2}^{n}+x_{3}^{n}+\ldots+x_{n}^{n}=\frac{n\left[1-(-1)^{n+1}\right]}{2(n+1)}
\end{array}\right.
$$  

Подставляя найденное выражение $c_{n}$, получим формулу Чебышева:

$$
\int_{-1}^{1} f(x) d x \approx \frac{2}{n}\left[f\left(x_{1}\right)+f\left(x_{2}\right)+\ldots f\left(x_{n}\right)\right]
$$

где точки $x_{1}, x_{2}, \ldots, x_{n}$ определяются из системы уравнений.   



# Квадратурная формула Гаусса.

В формулах численного интегрирования Ньютона-Котеса используются равноотстоящие узлы. В случае квадратурных формул Гаусса узлы интегрирования $x_{i}$ на отрезке $\left[x_{j-1}, x_{j}\right]$ располагаются не равномерно, а выбираются таким образом, чтобы при наименьшем возможном числе узлов точно интегрировать многочлены наивысшей возможной степени.

$$
\int_{a}^{b} f(x) \cdot d x \approx \frac{b-a}{2 N} \sum_{j=1}^{N} \sum_{i=0}^{n} c_{i n} \cdot f\left(x_{i}\right)
$$

Узлы $x_{i}$ являются корнями полинома Лежандра степени $n$, а веса вычисляются интегрированием полиномов Лежандра по формуле $a_{i}=\frac{2}{\left(1-x_{i}^{2}\right)\left[P_{n}^{\prime}\left(x_{i}\right)\right]^{2}}$, где $P_{n}^{\prime}-$ первая производная полинома Лежандра.  



\begin{aligned}
&\text { Таблица 2.2: Весовые коэффициенты и значения аргумента. }\\
&\begin{array}{cccc}
\hline \begin{array}{c}
\text { Число точек } n \\
\text { в квадратуре }
\end{array} & \begin{array}{c}
\text { Весовые } \\
\text { коэффициенты }
\end{array} & \begin{array}{c}
\text { Значения } \\
\text { аргумента }
\end{array} & \begin{array}{c}
\text { Погрешность } \\
\text { аппроксимации }
\end{array} \\
\hline 2 & w_{0}=1,000000000 & x_{0}=-0,577350269 & \sim f^{(4)}(x) \\
& w_{1}=1,000000000 & x_{1}=+0,577350269 & \\
3 & w_{0}=0,555555556 & x_{0}=-0,774596669 & \sim f^{(6)}(x) \\
& w_{1}=0,888888889 & x_{1}=+0,000000000 & \\
& w_{2}=0,555555556 & x_{2}=+0,774596669 & \\
4 & w_{0}=0,347854845 & x_{0}=-0,861136312 & \sim f^{(8)}(x) \\
& w_{1}=0,652145155 & x_{1}=-0,339981044 & \\
& w_{2}=0,652145155 & x_{2}=+0,339981044 & \\
& w_{3}=0,347854845 & x_{3}=+0,861136312 & \\
5 & w_{0}=0,236926885 & x_{0}=-0,906179846 & \sim f^{(10)}(x) \\
& w_{1}=0,478628670 & x_{1}=-0,538469310 & \\
& w_{2}=0,568888889 & x_{2}=+0,000000000 & \\
& w_{3}=0,478628670 & x_{3}=+0,538469310 & \\
& w_{4}=0,236926885 & x_{4}=+0,906179846 & \\
6 & w_{0}=0,171324492 & x_{0}=-0,932469514 & \sim f^{(12)}(x) \\
& w_{1}=0,360761573 & x_{1}=-0,661209386 & \\
& w_{2}=0,467913935 & x_{2}=-0,238619186 & \\
& w_{3}=0,467913935 & x_{3}=+0,238619186 & \\
& w_{4}=0,360761573 & x_{4}=+0,661209386 & \\
& w_{5}=0,171324492 & x_{5}=+0,932469514 & \\
\hline
\end{array}
\end{aligned}  



# Интегралы с особенностями.

Қак применять приведенные выше формулы для вычисления определенных интегралов с особенностями? Например, в случае, когда требуется вычислить интеграл

$$
\int_{0}^{1} \frac{f(x)}{\sqrt{x}} d x
$$

где $f(x)$ - гладкая на интервале интегрирования функция. Подынтегральная функция стремится к бесконечности при $x \rightarrow \infty$ и прямое применение вышеприведенных методов вычисления интеграла, как правило, невозможно.
В этом случае можно выделить особенность в нуле при помощи простого преобразования

$$
\int_{0}^{1} \frac{f(x)}{\sqrt{x}} d x=\int_{0}^{1} \frac{f(x)-\varphi(x)}{\sqrt{x}} d x+\int_{0}^{1} \frac{\varphi(x)}{\sqrt{x}} d x
$$

где функция $\varphi(x)$ выбирается так, чтобы первый интеграл в правой части не содержал особенности, а второй - вычислялся аналитически. Это справедливо, например, если в качестве $\varphi(x)$ взять отрезок разложения функции $f(x)$ в ряд Тейлора в точке $x=0 .$ Тогда

$$
\int_{0}^{1} \frac{f(x)-f(0)-x f^{\prime}(0)}{\sqrt{(} x)} d x+f(0) \int_{0}^{1} \frac{1}{\sqrt{x}} d x+f^{\prime}(0) \int_{0}^{1} \sqrt{x} d x .
$$  

Пример. Для $f(x)=\cos x$ вычисление интеграла $(2.9)$ сводится к вычислению

$$
\int_{0}^{1} \frac{\cos x}{\sqrt{x}} d x=\int_{0}^{1}(\cos x-1) x^{-1 / 2} d x+\int_{0}^{1} x^{-1 / 2} d x
$$

Второе слагаемое справа равно 2, а первое можно вычислить с помощью, например формулы Симпсона, которая дает значение $\approx 1,80905$ 




# Несобственные интегралы с внутренней особенностью.

Если особая точка $C \in[a, b]$, то используем простой приём, основанный на определении несобственного интеграла. Для этого интеграл представляют в виде: 

$\int_{a}^{b} f(x) d x=\int_{a}^{C-\delta_{1}} f(x) d x+\int_{C+\delta_{2}}^{b} f(x) d x+\int_{C-\delta_{1}}^{C+\delta_{2}} f(x) d x$,

$$
\int_{a}^{b} f(x) d x=\int_{a}^{C-\delta_{1}} f(x) d x+\int_{C+\delta_{2}}^{b} f(x) d x+\int_{C-\delta_{1}}^{C+\delta_{2}} f(x) d x
$$


причем $\delta_{1}, \delta_{2}$ выбирают столь малыми, чтобы в пределах заданной точности интеграл $\rho=\int_{C-\delta_{1}}^{C+\delta_{2}} f(x) d x$ не влиял бы на результат. 

Если вычисляется сходящийся несобственный интеграл 1-го рода 

$\int_{a}^{\infty} f(x) d x$, 

то для его приближенного вычисления используем равенство 

$\int_{a}^{\infty} f(x) d x=\int_{a}^{B} f(x) d x+\int_{B}^{\infty} f(x) d x .$ 

Причем число $B$ берут настолько большим, чтобы в пределах заданной точности интеграл $\int_{B}^{\infty} f(x) d x$ не влиял бы на результат. Далее последний интеграл вычисляют по какой-либо квадратурной формуле с нужной точностью.  




# Интегралы от быстроосциллирующих функций.

Вычисление интегралов от быстроосциллирующих функций встречается в задачах физики очень часто. Можно конечно вычислять интегралы от быстроосциллирующих функций, напрямую используя рассмотренные выше формулы, однако, это зачастую требует очень больших вычислительных затрат, поскольку шаг интегрирования нужно брать по крайней мере в несколько раз меньше периода осцилляций. Кроме того, вычисления должны проводиться с достаточной точностью, а при интегрировании осциллирующей функции происходит быстрая потеря точности из-за суммирования близких величин с противоположными знаками (интегралы от отдельных полупериодов ). Поэтому на практике используется другой прием.

Пример.Рассмотрим вычисление интеграла

$$
\int_{0}^{\pi} f(x) \sin k x d x
$$

где $k$ - большое число, например, 1000 . Аппроксимируем гладкую функцию $f(x)$ другой гладкой функцией $\varphi(x)$, для которой интеграл (2.10) вычисляется аналитически. Тогда задача вычисления интеграла от $f(x) \sin k x$ сводится к вычислению

$$
\int_{0}^{\pi} f(x) \sin k x=\int_{0}^{\pi} \varphi(x) \sin k x d x+\int_{0}^{\pi}[f(x)-\varphi(x)] \sin k x d x
$$

Если в качестве $\varphi(x)$ взять разложение $f(x)$ в ряд Тейлора, то второе слагаемое справа, как правило, мало и им можно пренебречь  

# Вычисление определенных интегралов (задачи).

## Задание 1. Двухточечная квадратура Гаусса.


Постройте квадратурную формулу Гаусса, интегрирующую точно многочлены степеней вплоть до третьей на интервале $[a, b]$. Заметим, что для этого достаточно построить _двухточечную_ квадратуру.

Напишите функцию, которая возвращает веса, $w_1$ и $w_2$, и узлы, $x_1$ и $x_2$, двухточечной квадратурной формулы Гаусса с весовой функцией $w(x) = 1$, т.е., интегралы вида

$$
\int_a^b\! f(x)\, dx \approx w_1 f(x_1) + w_2 f(x_2)
$$

<!-- grade_id: cell-cf3c018fa371396b, points: 0, grade: false, solution: true -->
```python
import numpy as np
def gauss_2(a, b):
    r"""Return nodes and weights for a two-point Gauss quadrature on [a, b].
    
    Parameters
    ----------
    a, b : float
       Boundaries of the integration interval
       
    Returns
    -------
    x1, x2, w1, w2 : float
       Nodes and weights of the quadrature.
    """
    x1 = ((b+a)/2 - (b+a)/(2*3**0.5))
    x2 = ((b+a)/2 + (b+a)/(2*3**0.5))
    w1 = (b-a)/2
    w2 = (b-a)/2
    return [x1,x2,w1,w2]
    raise NotImplementedError()
```

<!-- grade_id: cell-b6739a7f6045a9ec, points: 10, grade: true, solution: false -->
```python
from numpy.testing import assert_allclose

x1, x2, w1, w2 = gauss_2(0, 1)

def f(x, n): 
    return x**n

for n in [0, 1, 2, 3]:
    assert_allclose(w1*f(x1, n=n) + w2*f(x2, n=n),
                    1./(n+1), atol=1e-14)
```

## Задание 2. Вычитание интегрируемых сингулярностей.



Вычислите определённый интеграл методом трапеций с вычитанием сингулярности

$$
I = \int_{0}^{1}\frac{e^x}{\sqrt{x(1-x)}}dx.
$$

Вам могут пригодиться значения следующих определенных интегралов:

$$
\int_0^1 \frac{1}{\sqrt{x (1-x)}} \, dx=\pi,\quad \int_0^1 \frac{x}{\sqrt{x (1-x)}} \, dx=\pi/2.
$$

Преобразуйте данный интеграл, вычитая сингулярности. Выпишите расчетные формулы.

Cоставьте функцию, возвращающую значение интеграла методом трапеций.

<!-- grade_id: cell-7fb345ac8e79dced, points: 0, grade: false, solution: true -->
```python
def integ(npts=10):
    """Compute the value of the integral above.
    
    Subtract the singularities and use the trapezoid rule. 
    
    Parameters
    ----------
    npts : int
        The number of points for the trapezoid rule
        
    Returns
    -------
    I : float
       The computed value of the integral
    """
    x = np.linspace(0,1,npts)
    sum = np.zeros(npts-1)
    for i in range(npts-1):
        sum[i] = (x[i+1] - x[i])*(func(x[i+1]) + func(x[i]))/2
    
    return np.sum(sum) + np.pi/2 + 1/4*np.e*(2+np.pi)
    


def func(x):
    if (x==0) or (x==1):
        return 0
    elif (x<1/2):
        return (np.exp(x)-1)/((x*(1-x))**0.5)
    elif (x>=1/2):
        return np.e*(np.exp(x-1)-x)/((x*(1-x))**0.5)


integ(2000)
```

Output:
```
5.508427390879149
```

<!-- grade_id: cell-e24dd8dc92f2c8bd, points: 10, grade: true, solution: false -->
```python
# this is a test to check your computed value
from numpy.testing import assert_allclose

```

```python
assert_allclose(integ(2000), 5.50843, atol = 1e-4)
```

## Задание 3*. Интеграл от осциллирующей функции.



Рассмотрим определенный интеграл

$$
I = \int_0^\infty\! \frac{\sin(x) \cos{(\cos{(x)})}}{x}\,dx
$$

Вычислите значение данного интеграла с относительной точностью $10^{-6}$. Для упрощения задачи можете воспользоваться функционалом `scipy.integrate.quad`.

Заметим, что "из коробки" интеграл вычислить не удается, и нужно что-то придумать.

```python
from math import sin, cos

from scipy.integrate import quad
quad(lambda x: sin(x) * cos(cos(x)) / x, 0, float('inf'))
```

Output:
```
C:\Users\wowsl\AppData\Local\Temp\ipykernel_23680\778045123.py:4: IntegrationWarning: The integral is probably divergent, or slowly convergent.
  quad(lambda x: sin(x) * cos(cos(x)) / x, 0, float('inf'))
(1.9653912540956733, 4.089174284042322)
```

Напишите функцию, которая возвращает значение данного интеграла. Только само значение интеграла, без оценки погрешности.
(оборачиваем интеграл в функцию только для удобства автопроверки).

```python

def f(x):
    if (x==0):
        return 1
    else:
        return np.sin(x)*np.cos(np.cos(x))/x


def g(x):
    if (x==0):
        return 1
    else:
        return np.pi*np.sin(x)*np.cos(np.cos(x))/((x+np.pi)*x)


def h(x):
    if (x==0):
        return 1
    else:
        return 2*np.pi**2*np.sin(x)*np.cos(np.cos(x))/((x+2*np.pi)*(x+np.pi)*x)


```

<!-- grade_id: cell-ddd4e0b6d148fcfd, points: 0, grade: false, solution: true -->
```python
def A(n):
    b = np.pi
    x = np.linspace(0,b,n)
    sum = np.zeros(n-1)
    for i in range(n-1):
        sum[i] = (x[i+1] - x[i])*(f(x[i+1]) + f(x[i]))/2
    return np.sum(sum)
eps = 1e-7
start = 0
finish = 1
print(eps)
n = 200000

start = A(n-1)
finish = A(n)
n += 1
print(abs(finish - start))
    

print(finish)

```

Output:
```
1e-07
1.8053114558824745e-11
1.4373589464123515
```

```python
def B(n):
    b = np.pi
    x = np.linspace(0,b,n)
    sum = np.zeros(n-1)
    for i in range(n-1):
        sum[i] = (x[i+1] - x[i])*(g(x[i+1]) + g(x[i]))/2
    return np.sum(sum)
eps = 1e-7
start_2 = 0
finish_2 = 1
print(eps)
n = 200000
start_2 = B(n-1)
finish_2 = B(n)
print(abs(finish_2 - start_2))

    

print(finish_2)
```

Output:
```
1e-07
1.8052670469614895e-11
1.0737448897859745
```

```python
from scipy.integrate import quad
from math import sin, cos


def integ():
    """Return a single float, the computed value of the integral."""
    
    eps = 1e-6
    b = (16*np.pi**2/eps)**0.5
    n = int(b)*50
    x = np.linspace(0,b,n)
    sum = np.zeros(n-1)
    for i in range(n-1):
        sum[i] = (x[i+1] - x[i])*(h(x[i+1]) + h(x[i]))/2
    return np.sum(sum)/4 + finish/2 + finish_2/4 
    raise NotImplementedError()
integ()


```

Output:
```
1.203123852994621
```

<!-- grade_id: cell-b341c2ce36bc957d, points: 20, grade: true, solution: false -->
```python
from numpy.testing import assert_allclose

# this is a test cell, keep it intact

```

##  Задание 4.  Вычисление интеграла с внутренней особенностью.


Найти $\delta_{1}$ и $\delta_{2}$ такие, чтобы при вычислении интеграла
$$
J=\int_{-0,5}^{0,5}|x|^{-\frac{1}{2}} \cdot(1-x)^{-\frac{1}{2}} d x
$$
модуль $|\rho|<\varepsilon$, где $\varepsilon-$ требуемая точность расчетов.  

```python
import numpy as np
from scipy.integrate import quad
J = 2.88775422371971

def h(x):
    return 1/(abs(x)*(1-x))**0.5
        
eps = 1e-5
a = -1/2
b = 1/2
n = 1
x = np.linspace(a,b,n)
I = 1

while (eps <= abs(I - J)):
    right = b/n
    left = a/n
    second, eps_1 = quad(h,right,b)
    first, eps_2  = quad(h,a,left)
    I = first + second
    n*=2
    
    
print(abs(left))
print(abs(right))



```

Output:
```
2.9802322387695312e-08
2.9802322387695312e-08
```

## Задание 5.  Интеграл от быстроосциллирующей функции.



Вычислите интеграл от быстроосциллирующей функции
$$
\int_{0}^{\pi} \mathrm{e}^{-x} \sin k x d x
$$
для различных значений $k$ по методу Симпсона, а затем используя прием, описанный ранее, заменяя функцию $\mathrm{e}^{-x}$ ее интерполяционным полиномом второй степени (третьей степени). Совпадает-ли точность расчетов с ожидаемой? Сравните с точным значением интеграла, который равен

$$
\int_{0}^{\pi} \mathrm{e}^{-x} \sin k x d x=\frac{k}{1+k^{2}}-\frac{\mathrm{e}^{-\pi}(k \cos k \pi+\sin k \pi)}{1+k^{2}}
$$  

```python
def func(x,k):
    return np.exp(-x)*np.sin(k*x)

def J_true(k):
    return k/(1+k**2) - np.exp(-np.pi)*(k*np.cos(k*np.pi)+np.sin(k*np.pi))/(1+k**2)

def func_app(x,k):
    return  (1 - x + x**2/2 - x**3/6)*np.sin(k*x)
    

n = 10

sum = np.zeros(n-1)
a = 0
b = np.pi
k = np.linspace(1,100,10)
x = np.linspace(a,b,n)

for i in k:
    for j in range(n-1):
        sum[j] =((x[j+1]-x[j])/6*(func(x[j],i) + 4*func((x[j]+x[j+1])/2,i) + func(x[j+1],i)))
    print(np.sum(sum))


print("\n Approx values")
for i in k:
    for j in range(n-1):
        sum[j] =((x[j+1]-x[j])/6*(func_app(x[j],i) + 4*func_app((x[j]+x[j+1])/2,i) + func_app(x[j+1],i)))
    print(np.sum(sum))

print("\nTrue values")
for i in k:
    print(J_true(i))





```

Output:
```
0.5215961258838593
0.09449778627778296
-0.10446971655517477
-0.3827236016267149
0.11955785373948101
-0.14055896777695143
-0.20124761603190788
0.15623928858569636
0.18089110868871808
-0.12075422709765582

 Approx values
-0.23289982105970747
0.33360852069089897
0.13820869032552258
-1.2945917481530138
-0.15790759158866313
-0.477348675726222
0.2666207702950659
0.5517283327678587
-0.08690190492201859
-0.42629787097431027

True values
0.5216069591318861
0.07918229641954988
0.04527154739635238
0.028116444925697273
0.023171089003884378
0.01708001931056065
0.015566889203490588
0.012264472370653371
0.011720025085265809
0.009566904126949584
```

## Задание 6*.  Интегрирование с помощью ортогональных полиномов. 

Вычислите значение следующего интеграла:

$$
\int_{0}^{2 \pi} \frac{\cos (x) d x}{\sqrt{4 \pi^{2}-x^{2}}}
$$

используя соответствующий базис ортогональных полиномов. Найдите степень полинома (количество узлов), при которой точность превысит ошибку плавающей точки. Можете пользоваться вольфрамом или sympy.